# SLU04 - Basic Stats with Pandas: Exercises notebook

In these exercises we'll use a real dataset about house prices in King County in the US to learn how to obtain basic statistics from the data. This dataset comes from [Kaggle](https://www.kaggle.com/harlfoxem/housesalesprediction).


## Objectives

In these exercises the objective is for you to learn how to use pandas to obtain simple statistics of datasets.

## Dataset information

![](https://hookedonhouses.net/wp-content/uploads/2013/08/Capitol-Hill-Seattle-house-for-sale-E-Ward.jpg)

This dataset contains house sale prices for King County, which includes Seattle. It includes homes sold between May 2014 and May 2015.

The columns in the dataset are the following:

- `id`: a notation for a house
- `date`: Date house was sold
- `price`: Price is prediction target
- `bedrooms`: Number of Bedrooms/House
- `bathrooms`: Number of bathrooms/House
- `sqft_living`: square footage of the home
- `sqft_lot`: square footage of the lot
- `floors`: total floors in house
- `waterfront`: house which has a view to a waterfront
- `view`: has been viewed
- `condition`: how good the condition is (overall)
- `grade`: overall grade given to the housing unit, based on King County grading system
- `sqft_above`: square footage of house apart from basement
- `sqft_basements`: quare footage of the basement
- `yr_built`: built Year
- `yr_renovated`: Year when house was renovated
- `zipcode`: zip
- `lat`: latitude coordinate
- `long`: longitude coordinate
- `sqft_living15`: Living room area in 2015 (implies-- some renovations) This might or might not have affected the lotsize area
- `sqft_lot15`: lotSize area in 2015 (implies-- some renovations)

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import math 
import hashlib  # this is used to validate an answer, if you don't have it do !pip3 install hashlib
```

```python
data = pd.read_csv('data/kc_house_data.csv')
data.head()
```

Output:
```
id             date     price  bedrooms  bathrooms  sqft_living  \
0  7129300520  20141013T000000  221900.0         3       1.00         1180   
1  6414100192  20141209T000000  538000.0         3       2.25         2570   
2  5631500400  20150225T000000  180000.0         2       1.00          770   
3  2487200875  20141209T000000  604000.0         4       3.00         1960   
4  1954400510  20150218T000000  510000.0         3       2.00         1680   

   sqft_lot  floors  waterfront  view  ...  grade  sqft_above  sqft_basement  \
0      5650     1.0           0     0  ...      7        1180              0   
1      7242     2.0           0     0  ...      7        2170            400   
2     10000     1.0           0     0  ...      6         770              0   
3      5000     1.0           0     0  ...      7        1050            910   
4      8080     1.0           0     0  ...      8        1680              0   

   yr_built  yr_renovated  zipcode      lat     long  sqft_li
```

----

# Exercise 1 

Ok, let's start with very basic statistics.
 - How much is the most expensive house? What is its `id`?
 - How much is the cheapest house? What is its `id`?

<!-- grade_id: cell-493d5be6e9da4b41, points: 0, grade: false, solution: true -->
```python
data.set_index('id', drop=False)
price_most_expensive= data.price.max()
id_most_expensive= data.price.idxmax()
price_cheapest= data.price.min()
id_cheapest= data.price.idxmin()
```

##### Asserts 

<!-- grade_id: cell-681bb7a461c18fd9, points: 2, grade: true, solution: false -->
```python
assert price_most_expensive == 7700000.0
assert id_most_expensive == 6762700020
assert price_cheapest == 75000.0
assert id_cheapest == 3421079032
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-4-8ead28c4e653> in <module>
      1 assert price_most_expensive == 7700000.0
----> 2 assert id_most_expensive == 6762700020
      3 assert price_cheapest == 75000.0
      4 assert id_cheapest == 3421079032

AssertionError:
```

---

# Exercise 2 

Let's check the number of bedrooms.
- What is the maximum and minimum number of bedrooms in the dataset?
- What is the most common number of bedrooms?
- What is the average number of bedrooms?
- What is the median?
- What is the standard deviation?

<!-- grade_id: cell-3a582c9aad8a1b5e, points: 0, grade: false, solution: true -->
```python
# maximum = ...
# minimum = ...
# most_common = ...  # Hint: you should return a number, not a pandas Series :)
# average = ...
# median = ...
# standard_deviation = ...

# YOUR CODE HERE
raise NotImplementedError()
```

##### Asserts 

<!-- grade_id: cell-bf104255ad7aaf0c, points: 2, grade: true, solution: false -->
```python
assert maximum == 33
assert minimum == 0
assert most_common == 3
assert math.isclose(average, 3.37, abs_tol = 0.01)
assert median == 3
assert math.isclose(standard_deviation, 0.93, abs_tol = 0.01)
```

----

# Exercise 3 

In the previous exercise, you calculated some basic statistics for the number of bedrooms of the houses in the dataset. Have a look at the numbers you got again.
 - Are the minimum and maximum close to the mean and median?
 - Is the median smaller or larger than the mean? Did you expect that?
 - Given your answers to the questions above, do you expect the distribution of the number of bedrooms to be skewed? If so, do you expect it to be positively or negatively skewed?
 
Plot two histograms for the number of bedrooms, one with 10 bins and another with 30 bins. Can you visually confirm if the distribution is skewed?

Compute the skewness measure for the number of bedrooms. Is it positive or negative?

<!-- grade_id: cell-a7f823c7479df01a, points: 0, grade: false, solution: true -->
```python
## plot two histograms, side-by-side
# plt.figure(1, figsize=(12,4), dpi=200)
# plt.subplot(121)
## plot first histogram
# ...
# plt.subplot(122)
## plot first histogram
# ...

## compute skew measure
# skew = ...

# YOUR CODE HERE
raise NotImplementedError()
print("The skewness measure is {:.3f}.".format(skew))
```

##### Asserts 

<!-- grade_id: cell-e4f5228c7dbf23ad, points: 2, grade: true, solution: false -->
```python
assert math.isclose(skew, 1.974, abs_tol=0.001)
```

---

# Exercise 4 

Now let's look at the area of the homes and compare it to the area of the homes excluding the basement.
 - Do you expect the distributions of the areas to be skewed? If so, positively or negatively?
 - Which of the distributions do you expect to have more kurtosis (i.e., a longer "tail")? Why?
 
Verify the answers to the questions above by plotting a histogram (with 30 bins) for each of the distributions and by computing the skewness and kurtosis measures for each.

<!-- grade_id: cell-873fce2f25417734, points: 0, grade: false, solution: true -->
```python
# plt.figure(1, figsize=(12,4), dpi=200)
# plt.subplot(121)
# ...
# plt.subplot(122)
# ...
# skew_living = ...
# skew_above = ...
# kurt_living = ...
# kurt_above = ...

# YOUR CODE HERE
raise NotImplementedError()
print("The skewness measure for the area of the home distribution is {:.3f}.".format(skew_living))
print("The skewness measure for the area minus the basement distribution is {:.3f}.".format(skew_above))
print("The kurtosis measure for the area of the home distribution is {:.3f}.".format(kurt_living))
print("The kurtosis measure for the area minus the basement distribution is {:.3f}.".format(kurt_above))
```

##### Asserts 

<!-- grade_id: cell-d763ff473ff6ff99, points: 4, grade: true, solution: false -->
```python
assert(math.isclose(skew_living, 1.472, abs_tol=0.01))
assert(math.isclose(skew_above, 1.447, abs_tol=0.01))
assert(math.isclose(kurt_living, 5.243, abs_tol=0.01))
assert(math.isclose(kurt_above, 3.402, abs_tol=0.01))
```

---

# Exercise 5  

Find the quartiles of the area of the home and the area of the home minus basement. 

<!-- grade_id: cell-2449304f452c767a, points: 0, grade: false, solution: true -->
```python
# output quartiles as pandas Seroes
# living_quartiles = ...
# above_quartiles = ...

# YOUR CODE HERE
raise NotImplementedError()

display(pd.DataFrame({'Area of home Quantiles': living_quartiles, 
                      'Area minus basement Quantiles': above_quartiles}))
```

##### Asserts 

<!-- grade_id: cell-01d1b4a2cc6e9d83, points: 2, grade: true, solution: false -->
```python
assert (living_quartiles.sum() == 19717)
assert (above_quartiles.sum() == 14660)
assert list(living_quartiles.index) == [0.0, 0.25, 0.5, 0.75, 1.0]
```

---

# Exercise 6 

In this exercise, obtain:
* a list of all the years during which the houses were built;
* a list of all the years during which the houses were renovated.

*Reminder: none of the houses in this dataset was built and/or renovated when Jesus Christ was born.*

<!-- grade_id: cell-9e9944fd7252c53b, points: 0, grade: false, solution: true -->
```python
# list_years_built = ...
# list_years_renovated = ...

# YOUR CODE HERE
raise NotImplementedError()
```

##### Asserts 

<!-- grade_id: cell-7f879873e0ae96c8, points: 2, grade: true, solution: false -->
```python
assert isinstance(list_years_built, (list,))
assert len(list_years_built) == 116
assert isinstance(list_years_renovated, (list,))
assert len(list_years_renovated) == 69
```

----

# Exercise 7 

In Exercise 4, we saw that the area of the homes had a skewed distribution, with certainly a few outliers (*really huge houses!*).

Let's plot the histogram for the area again.

```python
data.sqft_living.plot.hist(bins=30, figsize=(10,6));
plt.xlim(0);
plt.xlabel('area of the home (in square feet)');
```

How to deal with the outliers?

In the Learning Notebook, you learned a few ways to deal with the outliers, in case they are negatively affecting your Machine Learning models. In this exercise, let's explore the **log transformation** and see if it helps us in this case.

Do the following:
* obtain the mean and the median of the areas of the homes; which one is greater?
* create a new field called `log_sqrt_living` with the log of the areas;
* obtain the mean and the median of the log of the areas; are they very different from each other?
* obtain the skewness measure for the distribution of the log of the areas; how does it compare to the one you obtained in Exercise 4?
* plot an histogram with 30 bins of the log of the areas.

What do you think? Were the outliers dealt with? How about the skewness of the distribution?

<!-- grade_id: cell-1956a03d1c17634e, points: 0, grade: false, solution: true -->
```python
## Compute mean and median of the areas of the homes
# area_mean = ...
# area_median = ...
## Create new field called `log_sqrt_living` with the log of the areas
# ...
## Compute mean and median of the log of the areas
# log_area_mean = ...
# log_area_median = ...
## Compute the skewness measure of the distribution of the log of the areas
# log_area_skew = ...
## Plot an histogram (with 30 bins) of the log of the areas
## For the x-axis, have the label 'log(area)'
# ...

# YOUR CODE HERE
raise NotImplementedError()
print('The home areas have mean %0.1f and median %0.1f' % (area_mean, area_median))
print('The log of areas have mean %0.3f and median %0.3f' % (log_area_mean, log_area_median))
print('The distribution of log of areas have a skewness measure of %0.3f' % log_area_skew)
```

##### Asserts 

<!-- grade_id: cell-065003e3c0d0a3e8, points: 6, grade: true, solution: false -->
```python
assert math.isclose(area_mean, 2079.9, abs_tol=0.1)
assert math.isclose(area_median, 1910.0, abs_tol=0.1)
assert math.isclose(log_area_mean, 7.55, abs_tol=0.01)
assert math.isclose(log_area_median, 7.55, abs_tol=0.01)
assert math.isclose(log_area_skew, -0.035, abs_tol=0.001)
assert math.isclose(data['log_sqft_living'].sum(), 163185.38, abs_tol=0.1)
```

----